---
id: 675433d8dc140478e1d73240
title: Add Friend Method
challengeType: 1
dashedName: add-friend-method
---

# --description--

**Introduction:**
Managing relationships between objects is crucial for creating interactive systems. This task involves enhancing an object with a method to maintain a list of friends.
<br>
**Challenge:**
Extend the createPersonObject function to include an addFriend method. This method should accept another person object and add it to a friends array property. If the friends array doesn’t exist, it should be created.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
