---
id: 675434137f447d7926cdaca5
title: Advanced Shoopping List
challengeType: 1
dashedName: advanced-shopping-list
---

# --description--

**Introduction:**
Enhancing a shopping list manager with methods to update and remove items adds greater functionality. This task involves expanding the basic shopping list with additional management capabilities.
<br>
**Challenge:**
Extend the createShoppingList function to include updateItemQuantity and removeItem methods. updateItemQuantity should accept an item name and update its quantity, while removeItem should accept an item name and remove it from the shoppingList.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
