---
id: 6754339ea79377788eeeeb02
title: Birthday Celebration Method
challengeType: 1
dashedName: birthday-celebration-method
---

# --description--

**Introduction:**
Enhancing object methods to manipulate properties dynamically is essential for interactive applications. This task involves adding a method to update an object's property in response to an event.
<br>

**Challenge:**
Extend the createPersonObject function to include a celebrateBirthday method. This method should increment the person’s age by 1 using this keyword to refer to the object's properties.

*Note:* 
Use `toUppercase()` in your code to change the alphabet to uppercase (in capital).

**Example:**
<br>
Input:
<br>
person : Object created by createPersonObject from Problem 1
<br>
age: 29
<br>
Output:
<br>
*Updated person object with incremented age.
<br>
person.celebrateBirthday();
<br>
console.log(person.introduce()); // "Hi, I'm John Doe, 30 years old, living at 987 Maple Street, New City."


# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
