---
id: 67543567604c727b38807d19
title: Caeser Cipher Encoder
challengeType: 1
dashedName: caeser-cipher-encoder
---

# --description--

**Introduction:**
The Caesar cipher is a classic encryption technique that shifts letters in the alphabet to encode messages. This task involves implementing a function to encrypt a string using this method.
<br>
**Challenge:**
Create the function caesarCipher(str, shift) that encrypts the input string by shifting each letter by the specified number of positions down the alphabet. The function should return the encrypted string.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
