---
id: 67543664e6f1637c38f64936
title: Calculating Distance Between Points
challengeType: 1
dashedName: calculating-distance-between-points
---

# --description--

**Introduction:**
Imagine you’re trying to measure the straight-line distance between two locations on a 2D map, like finding out how far apart two cities are. To calculate this distance, we use the distance formula, which involves both squaring and taking the square root of the differences between the coordinates of the two points. In JavaScript, Math.sqrt() helps us find the square root, and Math.pow() allows us to raise numbers to a power.
<br>
**Challenge:**
Your challenge is to write a function getDistance that takes four numbers (x1, y1, x2, y2) as arguments representing two points on a 2D plane. The function should return the distance between these points using the distance formula.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
