---
id: 675434ef1448267ad6ff9457
title: Capitalize Words
challengeType: 1
dashedName: capitalize-words
---

# --description--

**Introduction:**
Capitalizing the first letter of each word in a string enhances readability and presentation. This task involves creating a function to format strings into title case.
<br>
**Challenge:**
Create a function capitalizeWords(str) that capitalizes the first letter of each word in the input string and returns the updated string.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
