---
id: 6754369ae41dc17c5fbb756c
title: Compare Dice Results
challengeType: 1
dashedName: compare-dice-results
---

# --description--

**Introduction:**
Imagine you’re hosting a friendly dice-rolling competition between two players. Each player takes turns rolling a die, and you want to determine who wins each round based on the higher roll. To simulate this game in JavaScript, you need to generate random numbers between 1 and 6 to represent the die rolls. The game can continue for a set number of rounds or until one of the players reaches a target score. Use Math.random() to generate the random dice rolls and comparison operators to determine the winner of each round.
<br>
**Challenge:**
Your challenge is to create a JavaScript program that simulates a simple dice-rolling game for two players. Each player rolls a die (generate random numbers 1-6), and the player with the higher roll wins the round. The game continues for a certain number of rounds or until a player reaches a target score (Use Math.random, comparison operators).

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
