---
id: 675435e338706a7bc6911016
title: Custom Date Formatter
challengeType: 1
dashedName: custom-date-formatter
---

# --description--

**Introduction:**
The customDateFormatter function lets you format dates with specific placeholders, allowing for flexible and readable date representations. You can specify how you want the date and time components to appear in the output.
<br>
**Challenge:**
Implement the function customDateFormatter(date, formatString) to take a Date object and a format string with placeholders (YYYY, MM, DD, HH, mm, ss, A). Return a string where the date and time are formatted according to the given format.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
