---
id: 67543588baf80a7b6069c132
title: Date Adder
challengeType: 1
dashedName: date-adder
---

# --description--

**Introduction:**
Manipulating dates by adding days can be crucial for scheduling and time calculations. This task involves creating a function that adjusts a date by a given number of days.
<br>
**Challenge:**
Implement the function addDaysToDate(date, days) that takes a Date object and a number of days as input, and returns a new Date object representing the date after adding the specified number of days.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
