---
id: 675435c970bc547ba287b9d4
title: Days Until Birthday
challengeType: 1
dashedName: days-until-birthday
---

# --description--

**Introduction:**
Calculating the number of days until the next birthday helps with planning and reminders. This function will compute how many days are left until the next occurrence of a birthday.
<br>
**Challenge:**
Implement the function daysUntilBirthday(birthdayDateString) that takes a birthday date string in "YYYY-MM-DD" format and returns the number of days until the user's next birthday.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
