---
id: 67543473508572797156f17f
title: Detailed Shopping List
challengeType: 1
dashedName: detailed-shopping-list
---

# --description--

**Introduction:**
Displaying a shopping list with item details helps in better tracking and management. This task involves adding functionality to format and show a detailed list of items.
<br>
**Challenge:**
Extend the createShoppingList function to include a displayDetailedList method. This method should return a formatted string listing all items in the shoppingList along with their quantities.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
