---
id: 6754343d48db477949135b57
title: Enhanced Shopping List
challengeType: 1
dashedName: enhanced-shopping-list
---

# --description--

**Introduction:**
Adding functionality to check item existence and sort the list improves list management. This task involves extending a shopping list manager with these capabilities.
<br>
**Challenge:**
Extend the createShoppingList function to include checkItem and sortList methods. checkItem should accept an item name and return true if the item exists, while sortList should sort the shoppingList alphabetically by item name.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
