---
id: 67543370f9bb22786d6f0692
title: Enhanced Person Object
challengeType: 1
dashedName: enhanced-person-object
---

# --description--

**Introduction:**
Adding methods to update properties in objects enhances their flexibility and functionality. This task involves extending an existing function to include a method for updating an address.
<br>
**Challenge:**
Extend the createPersonObject function to include an updateAddress method. This method should accept a new address object and update the person’s address accordingly. Ensure the returned object includes this new method.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
