---
id: 67543315a0f4197822b69ef0
title: Flatten Nested Array
challengeType: 1
dashedName: flatten-nested-array
---

# --description--

**Introduction:**
Flattening a multi-dimensional array involves combining all nested elements into a single array. This task helps in practicing recursion and handling arrays of varying depths.
<br>
**Challenge:**
Write a recursive function flattenArray(arr) that takes a multi-dimensional array and returns a single flattened array with all elements combined. Ensure it handles arrays nested to any depth.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
