---
id: 675433bf0e3e1478af660590
title: List Person Details
challengeType: 1
dashedName: list-person-details
---

# --description--

**Introduction:**
Iterating over object properties to display their values is a useful skill for data representation and debugging. This task involves creating a function to list all details of a person object.
<br>
**Challenge:**
Write a function listPersonDetails that takes a person object (created by createPersonObject) and returns a string listing all properties and their values. Use a for-in loop to iterate over the object's properties.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
