---
id: 675435227cdba57af28d69ed
title: Palindrome Checker
challengeType: 1
dashedName: palindrome-checker
---

# --description--

**Introduction:**
Determining if a string is a palindrome is a common problem in text processing. This task involves creating a function that ignores spaces, special characters, and case to accurately check for palindromes.
<br>
**Challenge:**
Create a function isPalindrome(str) that checks if the input string is a palindrome, disregarding spaces, special characters, and case, and returns true if it is a palindrome and false otherwise.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
