---
id: 6754334ab3c1b278468adc18
title: Person Object Creator
challengeType: 1
dashedName: person-object-creator
---

# --description--

**Introduction:**
Creating structured objects with methods enhances code organization and functionality. This task involves building a function to create a person object with specific properties and methods.
<br>
**Challenge:**
Write a function createPersonObject that takes name, age, and address (with street and city properties) as parameters. The function should return an object with these properties and an introduce method that returns a string introduction of the person.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
