---
id: 675434b92364017a9764d618
title: Shopping List Enhancements
challengeType: 1
dashedName: shopping-list-enhancements
---

# --description--

**Introduction:**
Improving a shopping list with filtering and quantity calculations adds valuable insights for better inventory management. This task involves extending the list manager to include these functionalities.
<br>
**Challenge:**
Extend the createShoppingList function to include filterByQuantity and calculateTotalQuantity methods. filterByQuantity should return items with quantities greater than a given threshold, while calculateTotalQuantity should return the total quantity of all items in the shoppingList.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
