---
id: 675433f3c26c087905e22ced
title: Shopping List Manager
challengeType: 1
dashedName: shopping-list-manager
---

# --description--

**Introduction:**
Managing a shopping list involves adding and tracking items efficiently. This task focuses on creating a function to handle item addition to a list.
<br>
**Challenge:**
Write a function createShoppingList that initializes an empty shoppingList array. Implement an addItem function to accept an item object (with name and quantity properties) and add it to the shoppingList.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
