---
id: 675432e55ba97f76b33e2d3a
title: Simple Calculator Function
challengeType: 1
dashedName: simple-calculator-function
---

# --description--

**Introduction:**
Creating a basic calculator involves handling different operations based on user input. This task focuses on implementing a function that performs arithmetic calculations using operators.
<br>
**Challenge:**
Write an arrow function named calculate that takes two numbers and an operator (+, -, *, /) as arguments. Use a switch statement to perform the calculation and return the result. This exercise will help you practice using arrow functions and switch statements for conditional logic.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

**Example:**
<br>
Input:
<br>
let num1 = 10;
<br>
let num2 = 5
<br>
let Oper = +
<br>
Output:
<br>
15

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
