---
id: 67543537a33d657b1b9aa55f
title: String Compressor
challengeType: 1
dashedName: string-compressor
---

# --description--

**Introduction:**
Compressing a string by counting repeated characters can optimize storage and processing. This task involves creating a function to efficiently compress strings.
<br>
**Challenge:**
Implement the function compressString(str) that compresses the input string by replacing sequences of repeated characters with the character followed by its count, and returns the compressed string.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
