---
id: 675434d476f9b77ab6afecee
title: String Transformer
challengeType: 1
dashedName: string-transformer
---

# --description--

**Introduction:**
Transforming strings by removing spaces and replacing characters enhances data processing. This task focuses on updating a string by trimming and character substitution.
<br>
**Challenge:**
Create a function updatedString(str, originalChar, replacedChar) that removes leading and trailing spaces from the input string, replaces all occurrences of originalChar with replacedChar, and returns the updated string.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
