---
id: 675435ad51513d7b82408d5a
title: Weekend Checker
challengeType: 1
dashedName: weekend-checker
---

# --description--

**Introduction:**
Determining if a given date falls on a weekend is useful for scheduling and time management. This task involves creating a function to identify weekend dates.
<br>
**Challenge:**
Implement the function isWeekend(date) that takes a Date object as input and returns true if the date is a Saturday or Sunday; otherwise, return false.

# --instructions--

Click on this <a href = "https://cs50.ai/chat">Link</a>  to Go to CS50 AI 
And use this prompt prompt __________
Prompt: What does it mean to 'identify the number's sign' in this context?

# --hints--

You should use `typeof()`  in your code to check Data Type.

```js
assert(code.match(/typeof/g));
```

`checkDataType(67)` should return `number`

```js
assert(checkDataType(67)==="number")
```

`checkDataType("Hello")` should return `string`

```js
assert(checkDataType("Hello")==="string")
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js

```
